引数nの2乗を返す関数squareを定義します。「raise IntImplementedError()」を削除してからコードを記述してください。

<!-- grade_id: cell-31118fadd047a21d, points: 0, grade: false, solution: true -->
```python
def square(n):
    ### BEGIN SOLUTION
    return n*n
    ### END SOLUTION
```

このセルでプログラムの計算結果を確認しましょう。

```python
print(square(5))
```

Output:
```
25
```

<!-- grade_id: cell-331310679d2ed1e4, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert square(1) == 1
assert square(2) == 4
assert square(3) == 9
assert square(4) == 16
### END HIDDEN TESTS
```