# 関数の練習
今回は関数の練習です。


## 級数の和を求める関数
以下のsum関数は、引数として自然数nを受け取ると1からnまでの級数の和を返す関数です。

以下の「raise NotImplementedError()」を削除した後、その位置に適切なプログラムを記入してください。

<!-- grade_id: sum, points: 0, grade: false, solution: true -->
```python
def sum(n):
    if n < 1:
        return 0
    s = 0
    for i in range(1, n + 1, 1):
        ### BEGIN SOLUTION
        s = s + i
        ### END SOLUTION
    return s
```

以下のセルを利用して、sum関数の入出力を確認しましょう。

```python
# 例えば print(sum(4)) のような文を実行して表示内容を確認してみる
print(sum(4))
```

<!-- grade_id: correct_sum, points: 1, grade: true, solution: false -->
```python
# このセルには何も記入しないでください。
### BEGIN HIDDEN TESTS
assert sum(1) == 1
assert sum(4) == 10
assert sum(5) == 15
assert sum(0) == 0
assert sum(-1) == 0
### END HIDDEN TESTS
```

## 2乗を求める関数
以下のsquares関数は、引数として自然数nを受け取ると、1からnまでの各整数の2乗を求め、そのリスト（配列）を返します。

例えば、引数として10を渡すと、[1, 4, 9, 16, 25, 36, 49, 64, 81, 100] というリストを返します。

以下の「raise NotImplementedError()」を削除した後、その位置に適切なプログラムを記入してください。

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    if n < 1:
        raise ValueError("引数nには自然数を渡すべきです。")
    #### BEGIN SOLUTION
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

```python
# 例えば print(squares(10)) のような文を実行して表示内容を確認してみる
print(squares(10))
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
# このセルには何も記入しないでください。
### BEGIN HIDDEN TESTS
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
### END HIDDEN TESTS
```

<!-- grade_id: squares_invalid_input, points: 0, grade: true, solution: false -->
```python
# このセルには何も記入しないでください。
### BEGIN HIDDEN TESTS
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
### END HIDDEN TESTS
```